The optimal clutch size revisited: separating individual quality from the parental survival costs of reproduction

Author response:

The following is the authors’ response to the previous reviews.

Public Reviews:

Reviewer #4 (Public review):

We would like to thank the reviewer for their careful consideration of our manuscript. The suggestions have been useful in improving our manuscript. Please see our responses to the specific comments below.

Summary:

This is an important study that underscores that reproduction-survival trade-offs are not manifested (contrary to what generally accepted theory predicts) across a range of studies on birds. This has been studied by a meta-analytical approach, gathering data from a set of 46 papers (30 bird species). The overall conclusion is that there are no trade-offs apparent unless experimental manipulations push the natural variability to extreme values. In the wild, the general pattern for within-species variation is that birds with (naturally) larger clutches survive better.

Strengths:

I agree this study highlights important issues and provides good evidence of what it claims, using appropriate methods.

Weaknesses:

I also think, however, that it would benefit from broadening its horizon beyond bird studies. The conclusions can be reinforced through insights from other taxa. General reasoning is that there is positive pleiotropy (i.e. individuals vary in quality and therefore some are more fit (perform better) than others. Of course, this is within their current environment (biotic, abiotic, social. ...), with consequences of maintaining genetic variation across generations - outlined in Maklakov et al. 2015 (https://doi.org/10.1002/bies.201500025). This explains the outcomes of this study very well and would come to less controversy and surprise for a more general audience.

I have two fish examples in my mind where this trade-off is also discounted. Of course, given that it is beyond brood-caring birds, the wording in those studies is slightly different, but the evolutionary insight is the same. First, within species but across populations, Reznick et al. (2004, DOI: 10.1038/nature02936) demonstrated a positive correlation between reproduction and parental survival in guppies. Second, an annual killifish study (2021, DOI: 10.1111/1365-2656.13382) showed, within a population, a positive association between reproduction and (reproductive) aging.

In fruit flies, there is also a strong experimental study demonstrating the absence of reproduction-lifespan trade-offs (DOI: 10.1016/j.cub.2013.09.049).

I suggest that incorporating insights from those studies would broaden the scope and reach of the current manuscript.

We would like to thank the reviewer for this useful insight and for highlighting these studies. We have added detail in our discussion around positive correlations observed in the wild, and how positive pleiotropy has been presented as an explanation. We have also added the suggested studies as references to demonstrate the reproduction-lifespan trade-off has been shown to be absent. See lines 257-260.

Likely impact:

I think this is an important contribution to a slow shift in how we perceive the importance of trade-offs in ecology and evolution in general. While the current view still is that one individual excelling in one measure of its life history (i.e. receiving benefits) must struggle (i.e. pay costs) in another part. However, a positive correlation between all aspects of life history traits is possible within an individual (such as due to developmental conditions or fitting to a particular environment). Simply, some individuals can perform generally better (be of good quality than others).

We would like to thank the reviewer for highlighting the importance of our study. We hope our study will help the research community reflect on the importance of trade-offs between life-history traits and consider other possible explanations as to why variation in life-history traits is maintained within species.

Changes from the earlier version to the revised version

Abstract

Abstract

Life-history theory, central to our understanding of diversity in morphology, behaviour and senescence, describes how traits evolve through the optimisation of trade-offs in investment. Despite considerable study, there is only minimal support for trade-offs within species between the two traits most closely linked to fitness – reproductive effort and survival – questioning the theory’s general validity. We used a meta-analysis to separate the effects of individual quality (positive survival/reproduction correlation) from the costs of reproduction (negative survival/reproduction correlation) using studies of reproductive effort and parental survival in birds. Experimental enlargement of brood size caused reduced parental survival. However, the effect size of brood size manipulation was small and opposite to the effect of phenotypic quality, as we found that individuals that naturally produced larger clutches also survived better. The opposite effects on parental survival in experimental and observational studies of reproductive effort provides the first meta-analytic evidence for theory suggesting that quality differences mask trade-offs. Fitness projections using the overall effect size revealed that reproduction presented negligible costs, except when reproductive effort was forced beyond the maximum level observed within species, to that seen between species. We conclude that there is little support for the most fundamental life-history trade-off, between reproductive effort and survival, operating within a population. We suggest that within species, the fitness landscape of the reproduction–survival trade-off is flat until it reaches the boundaries of the between-species fast–slow life-history continuum. Our results provide a quantitative explanation as to why the costs of reproduction are not apparent and why variation in reproductive effort persists within species.

Added in the revised version: Introduction

Removed in the revised version: Main

Across taxa, we see wide variation in life-history traits, such as the number of offspring produced and time spent raising young 1–3. The central idea in life-history theory is that resources are finite, forcing trade-offs, meaning that investment in one aspect of life requires the sacrifice of another 4–7. As reproduction is one of the most resource-demanding life stages, it is expected that within-species variation in offspring production will be constrained by the cost of producing and raising young. It is thought that the fitness costs of reproduction are largely incurred as a decrement to parental survival, explaining why a species evolved a certain amount of reproductive effort, whilst also explaining the fast–slow life-history continuum between reproduction and lifespan observed across species 6. As reproduction and survival are the two components of life-history most closely related to fitness, this central trade-off has been the subject of much theoretical and empirical research, branching fields ranging from ecology, evolutionary biology to ageing research 8–10.

Brood size manipulations of birds in natural conditions have provided arguably the best experimental paradigm in which to test the survival costs of reproduction, as they directly alter parental care investment. Experimental increases in brood size result in increased parental effort, suggesting that parents can typically cope with increased reproductive demands 11–15. However, the expected increased costs of reproductive effort are not always detected and the current estimate of the cost to reproduce across studies suggests only a small and inconsistent effect 16–19. The absence of a cost of reproduction (determined through experimental studies) on survival means that costs must arise elsewhere or, alternatively, that individuals may differ in quality (determined through observational studies).

Earlier version: By quality, we refer to individuals operating at their own maximum reproductive output, determined by their phenotypic condition, local or temporal genetic adaptation, and the surrounding environment, and should manifest as a positive correlation between two fitness-associated traits (such as reproduction and survival) 11, 19–22. The relative importance of the trade-off between reproductive effort and survival to life-history theory and the biology of ageing 9, therefore remains unclear. In addition, the compelling theoretical explanation for the lack of an apparent trade-off due to the confounding effects of individual quality has not been investigated on a quantitative level 5, 19, 23, 24. Our study aims to separate the effects of a trade-off between reproductive effort and parental survival and the effects of individual quality by comparing studies of brood manipulations and naturally varying clutch size.

Revised version: By quality, we refer to individuals operating at their own maximum reproductive output, determined by their phenotypic condition, local or temporal genetic adaptation, and the surrounding environment, and should manifest as a positive correlation between two fitness-associated traits (such as reproduction and survival) 11,19–22. The relative importance of the trade-off between reproductive effort and survival to life-history theory and the biology of ageing 9, therefore remains unclear. In addition, the compelling theoretical explanation for the lack of an apparent trade-off due to the confounding effects of individual quality has not often been investigated on a quantitative level 5,19,23,24. Our study aims to separate the effects of a trade-off between reproductive effort and parental survival and the effects of individual quality by comparing studies of brood manipulations and naturally varying clutch size.

Earlier version: Here, we present a meta-analysis that distinguishes between individual quality effects (naturally varying clutch size) and the costs of reproduction (brood manipulation studies). To do this, we tested how parental annual survival in birds is affected by the clutch size (note brood size and clutch size are used interchangeably throughout this manuscript) they cared for in two different contexts: first in brood manipulation studies and, second, in observational studies of natural variation in clutch size. We expressed changes in survival on a per-egg basis, which, for the first time, allows for a quantitative comparison across studies and thus species. We find the effect of quality is associated with higher survival chances given an increase in reproductive effort, and that this effect is opposite but equal in magnitude to the costs of reproduction. The survival trade-off for offspring production within a population is therefore offset by differences in quality, potentially constraining the evolution of higher reproductive effort. Our analysis also uniquely allowed a quantitative across species comparison, as survival risk was expressed on a per-egg basis. We transformed the response variable, scaling for variance and mean, given that a per-egg increase in clutch size does not equate to the same proportional increase in reproductive effort for all species equally. Our findings suggest that species that generally lay smaller clutches are affected more severely by brood-size manipulations. This provides evidence that trade-offs are only detected when an individual is forced to perform far outside its optimum level (Figure S1). Species who on average have smaller clutch sizes receive disproportionate manipulations (Figure S1A) and this pushes species beyond their natural range (Figure S1B).

Revised version: Here, we present a meta-analysis that distinguishes between individual quality effects and the costs of reproduction. To do this, we tested how parental annual survival in birds is affected by the clutch size they cared for using brood manipulation studies (costs of reproduction) and, observational studies of natural variation in clutch size (individual quality effects) (note brood size and clutch size are used interchangeably throughout this manuscript). We find the effect of quality is associated with higher survival chances given an increase in reproductive effort, and that this effect is opposite but equal in magnitude to the costs of reproduction. Our analysis also allowed us to transform the response variable, scaling for variance and mean, given that a per-egg increase in clutch size does not equate to the same proportional increase in reproductive effort for all species equally. Our findings suggest that species that generally lay smaller clutches are affected more severely by brood-size manipulations.

Earlier version: To predict the evolutionary consequences of the effect sizes that we estimated using the meta-analysis, we projected the fitness consequences for a change in clutch-size life-history strategy. We found that the costs to parental survival as estimated from brood size manipulation studies translate into negligible fitness costs, with a relatively flat fitness landscape, suggesting that birds underproduce in terms of clutch size, given the absence of fitness costs. This conclusion fits with our comparative finding that suggests that only brood sizes manipulated beyond the natural range incur substantial survival costs. Our results therefore suggest that, in wild populations, parental survival costs are, at most, a small component of the total fitness costs of investing parental effort. Our results do suggest that a cost of reproduction can be detected when an organism is pushed to the extreme of its reproductive capacity. We therefore infer that, though the survival–parental care trade-off does exist within species, it not strong enough to constrain clutch size and can therefore not explain clutch size evolution and thus variation within species. In addition, our work shows that differences in individual quality counterbalance the trade-off between survival and reproduction, as previously theorised 5, and as such could constrain reproductive effort further, and contribute to the maintenance of clutch size variation in a population.

Revised version: To predict the evolutionary consequences of the effect sizes that we estimated using the meta-analysis, we projected the fitness consequences for a change in clutch-size life-history strategy. We found that the costs to parental survival as estimated from brood size manipulation studies translate into negligible fitness costs, with a relatively flat fitness landscape, suggesting that birds underproduce in terms of clutch size, given the absence of fitness costs. Our results therefore suggest that, in wild populations, parental survival costs are, at most, a small component of the total fitness costs of investing parental effort. We therefore infer that, though the survival–parental care trade-off does exist within species, it not strong enough to constrain clutch size and can therefore not explain clutch size evolution and thus variation within species. 5

Results

Earlier version: The relationship between clutch size and survival was significantly different and opposite between observational and brood manipulation studies, irrespective of how clutch size was scaled (p < 0.01, Figure 1, Table 1). Within natural variation (observational studies), parents with larger clutches showed increased survival (or parents with smaller clutches showed reduced survival). In contrast, when broods were experimentally manipulated, the opposite relationship was found; increasing brood sizes decreased survival. Although the difference in overall effect size between experimental and natural variation in clutch size was strongly significant in each comparison made, the individual overall effect sizes only became significant (from zero) when clutch size was expressed as a proportional increase. Expressing clutch size as a proportional increase corrects for the variation in average clutch size observed across the species included in this analysis, which ranged from 2 to 11. The parental effort required to raise two instead of one chick is potentially doubled, whereas one additional chick in a brood of 11 is likely to require only a marginal increase in effort. Indeed, also when using a between-species comparison, the effects of clutch size manipulation and quality were strongest in the species that laid the smallest clutches, suggesting that costs to survival were only observed when a species was pushed beyond its natural limits (Figure 2, Figure S1A, Table S1).

Revised version: The relationship between clutch size and survival was significantly different and opposite between observational and brood manipulation studies, irrespective of how clutch size was scaled (p < 0.01, Figure 1, Table 1). Within natural variation (observational studies), parents with larger clutches showed increased survival (or parents with smaller clutches showed reduced survival). In contrast, when broods were experimentally manipulated, the opposite relationship was found; increasing brood sizes decreased survival. Although the difference in overall effect size between experimental and natural variation in clutch size was strongly significant in each comparison made, the individual overall effect sizes only became significantly different from zero when clutch size was expressed as a proportional increase. Expressing clutch size as a proportional increase corrects for the variation in average clutch size observed across the species included in this analysis, which ranged from 2 to 11. The parental effort required to raise two instead of one chick is potentially doubled, whereas one additional chick in a brood of 11 is likely to require only a marginal increase in effort. Indeed, also when using a between-species comparison, the effects of clutch size manipulation and quality were strongest in the species that laid the smallest clutches, suggesting that costs to survival were only observed when a species was pushed beyond its natural limits (Figure 2, Figure S1A, Table S1).

Earlier version: The effects size (log odds of survival given an increase in clutch size) for three different measures of clutch size: (A) raw, (B) standardised, and (C) proportional clutch size. Coloured points are the combined effect sizes of the odds ratios with their 95% confidence intervals. Points are coloured by whether they represent brood manipulation experiments (costs of reproduction) or they are observational (quality). Grey points are the odds ratios of each study, with their sizes weighted by the sampling variance.

Revised version: The effects size (log odds of survival given an increase in clutch size) for three different measures of clutch size: (A) raw, (B) standardised, and (C) proportional clutch size.

Added in the revised version: Coloured points are the combined effect sizes of the odds ratios with their 95% confidence intervals. Points are coloured by whether they represent brood manipulation experiments (costs of reproduction) or they are observational (quality). Grey points are the odds ratios of each study, with their sizes weighted by the sampling variance.

Added in the revised version: The meta-analytic linear regression (Table S1) of the effect size for increasing clutch size (per egg) on parental survival, given the average clutch size for the species.

Earlier version: The meta-analytic linear regression (Table S1) of the effect size for increasing clutch size (per egg) on parental survival, given the average clutch size for the species. Species with small clutch sizes showed stronger costs of reproduction and a stronger relationship with quality (Interaction between treatment and species clutch size effect size = −0.036, p = 0.015). The points are the survival effect sizes (log odds ratio) per egg (as in Figure 1A) on parental survival in each study, with the point size reflecting the meta-analytic weight of that study.

Revised version: Species with small clutch sizes showed stronger costs of reproduction and a stronger relationship with quality (Interaction between treatment and species clutch size effect size = -0.036, p = 0.015). The points are the survival effect sizes (log odds ratio) per egg (as in Figure 1A) on parental survival in each study, with the point size reflecting the meta-analytic weight of that study.

Earlier version: Males and females did not differ in their survival response to changing clutch size (Table S2, Figure S2, contrary to Santos & Nakagawa 2012 18). The variance assigned to the random effects in the model was largely accounted for by study (Table S3). Species accounted for more variation than the phylogeny, indicating that species vary in the relationship between survival and reproductive effort, irrespective of their shared evolutionary ancestry. However, our dataset included few closely related species, which reduces our ability to estimate phylogenetic effects (Figure 3).

Revised version: Males and females did not differ in their survival response to changing clutch size (Table S2, Figure S2, contrary to Santos & Nakagawa 2012 18). The variance assigned to the random effects in the model was largely accounted for by study (Table S3). Species identity accounted for more variation than the phylogeny, indicating that species vary in the relationship between survival and reproductive effort, irrespective of their shared evolutionary ancestry. However, our dataset included few closely related species, which reduces our ability to estimate phylogenetic effects (Figure 3).

Added in the revised version: Phylogenetic tree of species included in our meta-analysis.

Earlier version: Phylogenetic tree of species included in our meta-analysis. Numeric labels denote the node support (as a percentage). Coloured points indicate whether the species were included in the experimental or observational studies.

Revised version: Numeric labels denote the node support (as a percentage). Coloured points indicate whether the species were included in the experimental or observational studies.

Projected fitness consequences of the costs of reproduction

Earlier version: From our meta-analysis we now have a quantifiable and comparable effect size for the survival costs of reproduction that we can use to predict the evolutionary consequences of individuals in a given population increasing their clutch size across a range of life histories. To this end, we projected the fitness consequences of increased reproductive effort, starting with the average effect size estimate per egg (−0.05, Table 1) across a range of life histories, for a range of annual survival rates and clutch sizes (Figure 4). Overall, the effect size estimated in the meta-analysis (−0.05) resulted in a gain of fitness when reproductive output increased, especially in hypothetical species with low survival and small clutches. If we used a more severe effect size (−0.15, the lower confidence interval of our meta-analysis), this led to a reduction in fitness in almost all cases. Conversely, the benefit of higher reproductive output was largely offset by the cost of survival when a species’ survival rate and clutch size were high. When we increased the effect size up to five-fold, fitness costs of reproduction became more pronounced, but were still not present in species with small clutches and short lifespans.

Revised version: From our meta-analysis we now have a quantifiable and comparable effect size for the survival costs of reproduction that we can use to predict the evolutionary consequences of individuals in a given population increasing their clutch size across a range of life histories. To this end, we projected the fitness consequences of increased reproductive effort, starting with the average effect size estimate per egg (-0.05, Table 1) across a range of life histories, for a range of annual survival rates and clutch sizes (Figure 4). Overall, the effect size estimated in the meta-analysis (-0.05) resulted in a gain of fitness when reproductive output increased, especially in hypothetical species with low survival and small clutches. If we used a stronger effect size (-0.15, the lower confidence interval of our meta-analysis), this led to a reduction in fitness in almost all cases. Conversely, the benefit of higher reproductive output was largely offset by the cost of survival when a species’ survival rate and clutch size were high. When we increased the effect size up to five-fold, fitness costs of reproduction became more pronounced, but were still not present in species with small clutches and short lifespans.

Added in the revised version: Isoclines of selection differentials among hypothetical control populations (in which individuals reproduce at the species’ mean rate) and hypothetical brood-manipulated populations (where individuals reproduce at an increased rate compared to control) for their whole lives.

Earlier version: Isoclines of selection differentials among hypothetical control populations (in which individuals reproduce at the species’ mean rate) and hypothetical brood-manipulated populations (where individuals reproduce at an increased rate compared to control) for their whole lives. Selection differentials (i.e., the relative difference in lifetime reproductive output between hypothetical control and brood-manipulated populations) above 1 represents high lifetime fitness. Survival rates, clutch sizes, the magnitude of the manipulation (chicks added) and effect sizes represent the range of these variables present in the studies used in our meta-analysis. For each clutch size, we used a predicted survival rate and effect size to give isoclines that are biologically meaningful (exemplar birds shown in red). Arrows indicate the relative size and direction of selection in life-history space (on the reproduction axis). The costs of reproduction we estimated within species are predicted to result in a fast–slow life-history continuum across species, and the exemplar species we used as examples fit on this diagonal of survival rate/ clutch size combinations. We suggest that individual species show limited costs of reproduction, as they operate within relatively wide constraints imposed by the cost of reproduction that is responsible for the strong life-history trade-off observed across species.

Revised version: Selection differentials (i.e., the relative difference in lifetime reproductive output between hypothetical control and brood-manipulated populations) above 1 represents high lifetime fitness. Survival rates, clutch sizes, the magnitude of the manipulation (chicks added) and effect sizes represent the range of these variables present in the studies used in our meta-analysis. For each clutch size, we used a predicted survival rate and effect size to give isoclines that are biologically meaningful (exemplar birds shown in red). Arrows indicate the relative size and direction of selection in life-history space (on the reproduction axis). The costs of reproduction we estimated within species are predicted to result in a fast–slow life-history continuum across species, and the exemplar species we used as examples fit on this diagonal of survival rate/ clutch size combinations. We suggest that individual species show limited costs of reproduction, as they operate within relatively wide constraints imposed by the cost of reproduction that is responsible for the strong life-history trade-off observed across species.

Under long-term evolution these selection differentials should lead to individual hypothetical species moving towards the diagonal (bottom left to the right top corner). This diagonal represents the observed fast–slow pace of life continuum observed among species 25. Exemplar species (i.e., with survival and average clutch size combinations observable in wild populations), for which we predicted the fitness consequences of the costs of reproduction, lie on this comparative diagonal in life history. In these exemplar species, the selection differential was observed to lie slightly above one, indicating that individuals having a higher clutch size than the species’ average would gain a slight fitness benefit. The fitness costs and benefits did, in general, not diverge substantially with the addition of chicks, but flattened, suggesting that the costs of survival counterbalance the benefits of reproduction across a range of reproductive outputs within a species.

The low costs of reproduction that we estimated could still be responsible for between-species life-history evolution, constraining species reproductive output and survival combination to fall along the diagonal of the fast–slow pace of life continuum. How selection pressures translate into short term and longer-term evolutionary trajectories is uncertain. Often directional selection estimated in the wild does not translate to the inter-generational change on the population level 26. Note, however, that only far away from the diagonal did our fitness projections reach a magnitude that would be predicted to lead to rapid evolutionary change 27 and see SI). The weak selection effects that lie on the diagonal are probably to be counterbalanced in the wild by factors such as environmental effects and genetic effects (e.g., gene flow from immigration or random mutation) 28. We argue that within-species the minimal costs of reproduction, a flat fitness landscape and quality effects (Figure 1) together explain why individuals appear to under-produce. Only when individuals are pushed beyond the observed between-species constraint do costs become apparent (Figures 1, 2).

Our interpretation of the reproduction/lifespan life-history trade-off, based on our quantitative meta-analysis and subsequent fitness projections, explains several key observations and contradictions in the field. A strong trade-off is observed between species, but within species this trade-off is not apparent and variation in reproductive output is maintained within fitness boundaries similar to those that determine the between-species life history trade-off. The implication of this conclusion is that the costs of reproduction are likely to operate on a physiological level, but that the fitness consequences will remain largely flat over a species’ observed variation in reproductive output. These effects are further obscured by the effects of quality (observational studies), which are opposite in sign and magnitude to the cost of reproduction (experimental studies) (Figure 1) and are likely to further flatten the fitness landscape.

Discussion

Earlier version: Our results provide the first meta-analytic evidence that differences in individual quality drive variation in clutch size and that parental survival costs do not constrain within-species variation in reproductive effort (Figure 5). Here, we use the definition of quality as a combination of traits that give an individual higher fitness 5. The finding of individual quality being a driver of variation in clutch size is not necessarily a surprising one, even though this contradicts the general theory that trade-offs drive variation in clutch size, as among-individual heterogeneity is also a well-established effect across populations in the field of ecology and evolution (such as age-specific reproductive performance) 5, 24. Furthermore, a recent meta-analysis on genetic life-history trade-offs found several positive genetic correlations between survival and other life-history traits, which is in conflict with life-history trade-off theory 29. Brood manipulations often fail to take into account that variation in resource acquisition may be larger than variation in resource allocation, and though intuitively manipulation of effort tests variation in resource allocation, it fails to account for the baseline variation in reproductive performance 5. This presents a fundamental flaw which our results demonstrate – variation in resource allocation can be present and testable through experiments (i.e., brood manipulations) but may only account for a small proportion of variation observed in reproductive output. Perhaps fitting with this interpretation or with additional costs from recurring bouts of enhanced reproductive effort, Jackdaws (Corvus monedula) show costs of increased brood size only when applied sequentially across years, but note that for this study there is no control group that is only manipulated in single years30.

Revised version: Our results provide the first meta-analytic evidence that differences in individual quality drives variation in clutch size and that parental survival costs do not constrain within-species variation in reproductive effort (Figure 5). Here, we use the definition of quality as a combination of traits that give an individual higher fitness 5. The finding of individual quality being a driver of variation in clutch size is not necessarily a surprising one, even though this contradicts the general theory that trade-offs drive variation in clutch size, as among-individual heterogeneity is also a well-established effect across populations in the field of ecology and evolution (such as age-specific reproductive performance) 5,24. There has previously been discussion of why positive genetic correlations have been found in the wild and result from, for example, positive pleiotropy 29. Furthermore, empirical data has also found an absence of a trade-off between reproduction and survival 30–32. A recent meta-analysis on genetic life-history trade-offs found several positive genetic correlations between survival and other life-history traits, which is in conflict with life-history trade-off theory 33. Brood manipulations often fail to take into account that variation in resource acquisition may be larger than variation in resource allocation, and though intuitively manipulation of effort tests variation in resource allocation, it fails to account for the baseline variation in reproductive performance 5. This presents a fundamental flaw which our results demonstrate – variation in resource allocation can be present and testable through experiments (i.e., brood manipulations) but may only account for a small proportion of variation observed in reproductive output. Perhaps fitting with this interpretation or with additional costs from recurring bouts of enhanced reproductive effort, Jackdaws (Corvus monedula) show costs of increased brood size only when applied sequentially across years, but note that for this study there is no control group that is only manipulated in single years 34.

Decision tree representing the logical steps from our original hypothesis to our overall interpretation of our findings.

Earlier version: The reason selection is not acting on “high quality” individuals is currently unknown, but it is likely that environmental variability leads to alternative phenotypes being selected for at different points in space and time (also discussed in Pujol et al., 2018 26). It is also possible that the quality effect could represent a terminal effect, where individuals have lower reproductive output in the year preceding their death, thereby driving the trend for naturally lower laying birds to have lower survival when estimated between individuals (e.g., Coulson & Fairweather, 2001; Rattiste, 2004; Simons et al., 2016 31–33; also see Hammers et al., 2012 34 for age-related changes in reproductive output). The effect on parental annual survival of having naturally larger clutches was significantly opposite to the result of increasing clutch size through brood manipulation, and quantitatively similar. Parents with naturally larger clutches are thus expected to live longer and this counterbalances the “cost of reproduction” when their brood size is experimentally manipulated. It is, therefore, possible that quality effects mask trade-offs. Furthermore, it could be possible that individuals that lay larger clutches have smaller costs of reproduction, i.e. would respond less in terms of annual survival to a brood size manipulation, but with our current dataset we cannot address this hypothesis (Figure 5). The effect of a change in clutch size on parental survival may also be non-linear, but it remains to be determined what shape, if not linear, the relationship would be. Although these possible non-linear effects warrants investigation, these relationships likely differ between species and inclusion in our current work here could lead to spurious relationships being reported.

Revised version: The reason selection is not acting on “high quality” individuals is currently unknown, but it is likely that environmental variability leads to alternative phenotypes being selected for at different points in space and time (also discussed in Pujol et al., 2018 26). It is also possible that the quality effect could represent a terminal effect, where individuals have lower reproductive output in the year preceding their death, thereby driving the trend for naturally lower laying birds to have lower survival when estimated between individuals (e.g., Coulson & Fairweather, 2001; Rattiste, 2004; Simons et al., 2016 35–37; also see Hammers et al., 2012 38 for age-related changes in reproductive output). The effect on parental annual survival of having naturally larger clutches was significantly opposite to the result of increasing clutch size through brood manipulation, and quantitatively similar. Parents with naturally larger clutches are thus expected to live longer and this counterbalances the “cost of reproduction” when their brood size is experimentally manipulated. It is, therefore, possible that quality effects mask trade-offs. Furthermore, it could be possible that individuals that lay larger clutches have smaller costs of reproduction, i.e. would respond less in terms of annual survival to a brood size manipulation, but with our current dataset we cannot address this hypothesis (Figure 5). The effect of a change in clutch size on parental survival may also be non-linear, but it remains to be determined what shape, if not linear, the relationship would be. Although these possible non-linear effects warrants investigation, these relationships likely differ between species and inclusion in our current work here could lead to spurious relationships being reported.

For both costs of reproduction and quality effects, we found that species that laid the smallest clutches showed the largest effects. Brood manipulations that affect parental survival are thus likely to be the result of pushing parental effort beyond its natural optima (Figure S1) and so arguably that brood manipulations are not necessarily a good test of whether trade-offs happen in the wild. The classic trade-off between adult survival and the clutch size cared for is only apparent when an individual is forced to raise a clutch outside of its individual optimum, and these effects are confounded or even fully counterbalanced by differences in quality (as theorised in van Noordwijk & de Jong, 1986 5).

Earlier version: Our fitness projections of the consequences of the costs of reproduction using the overall effect size we estimated suggest that, for current extant species, the within-species fitness landscape of the reproduction survival trade-off is flat. Species’ life history decisions are constrained within a broader fast–slow life-history continuum, explaining why variation within species in reproductive effort, such as in clutch size, is large and near universal. Of course, this analysis does not fully explore more complex variation in clutch size observed in the wild, such as that observed between temperate and tropical reproductive strategies. Such work would provide useful in understanding why variation in clutch size is maintained across species with different life-history strategies. Our results do, however, provide a key step in understanding the generalised relationship between reproductive effort and survival across species (Figure 5). Our interpretation also assumes that other fitness costs of reproduction are smaller or at least less relevant than survival costs. However, it is possible that such costs are actually more important, and it should be noted that effects such as those on offspring quality (e.g., Conrad & Robertson, 1992 and Smith et al., 1989 14, 35), parental condition other than survival (e.g., Reid, 1987 and Kalmback et al., 2004 36, 37) and future reproductive effort (e.g., Järvistö et al., 2016 38) have been observed (also see Figure 5). However, the importance of these effects is likely to vary considerably in different species. Using offspring quality as an example, some species produce sacrificial offspring, others experience catch-up growth, meaning that though the effect of increasing offspring number on offspring quality is important for certain species, drawing generalised across-species conclusions is unlikely to be possible. Conversely, the survival cost to reproduce is thought to be universal across both bird species and across taxa, and it is perhaps for this reason that the reproduction-survival trade-off has been considered to contribute more to variation in reproductive effort than other trade-offs. Our work refutes this idea directly. Indeed, the few studies that have measured the different domains that contribute to fitness in brood-size manipulation studies concluded that only in combination do these costs result in balancing selection for the current most common brood size in the population 39–41. Such classic trade-off explanations do, however, fail to explain why variation in reproductive effort is prevalent within species and why between-species life-history trade-offs appear so much stronger and conclusive. Our analysis and interpretation suggest that, at its optimum, the within-species trade-off between survival and reproduction is relatively flat, and thus neutral to selection (supporting the theory presented in Cohen et al., 2020 19). We suggest that the lack of evidence supporting trade-offs driving within-species variation does not necessarily mean that physiological costs of reproduction are non-existent (e.g., Smith et al., 1989 and Lemaître et al., 2015 35, 42), but rather that, within the wild and within the natural range of reproductive activities, such costs are not relevant to fitness. One key explanation for this effect supported by our meta-analysis (Figure 5) and prior theory 5, is that individuals differ in quality.

Revised version: Our fitness projections of the consequences of the costs of reproduction using the overall effect size we estimated suggest that, for current extant species, the within-species fitness landscape of the reproduction survival trade-off is flat. Species’ life history decisions are constrained within a broader fast–slow life-history continuum, explaining why variation within species in reproductive effort, such as in clutch size, is large and near universal. Of course, this analysis does not fully explore more complex variation in clutch size observed in the wild, such as that observed between temperate and tropical reproductive strategies. Such work would provide useful in understanding why variation in clutch size is maintained across species with different life-history strategies. Our results do, however, provide a key step in understanding the generalised relationship between reproductive effort and survival across species (Figure 5). Our interpretation also assumes that other fitness costs of reproduction are smaller or at least less relevant than survival costs. However, it is possible that such costs are actually more important, and it should be noted that effects such as those on offspring quality (e.g., Conrad & Robertson, 1992 and Smith et al., 1989 14,39), parental condition other than survival (e.g., Reid, 1987 and Kalmback et al., 2004 40,41) and future reproductive effort (e.g., Järvistö et al., 2016 42) have been observed (also see Figure 5). However, the importance of these effects is likely to vary considerably in different species. Using offspring quality as an example, some species produce sacrificial offspring, others experience catch-up growth, meaning that though the effect of increasing offspring number on offspring quality is important for certain species, drawing generalised across-species conclusions is unlikely to be possible. Conversely, the survival cost to reproduce is thought to be universal across both bird species and across taxa, and it is perhaps for this reason that the reproduction-survival trade-off has been considered to contribute more to variation in reproductive effort than other trade-offs. Our work refutes this idea directly. Indeed, the few studies that have measured the different domains that contribute to fitness in brood-size manipulation studies concluded that only in combination do these costs result in balancing selection for the current most common brood size in the population 43–45. Such classic trade-off explanations do, however, fail to explain why variation in reproductive effort is prevalent within species and why between-species life-history trade-offs appear so much stronger and conclusive. Our analysis and interpretation suggest that, at its optimum, the within-species trade-off between survival and reproduction is relatively flat, and thus neutral to selection (supporting the theory presented in Cohen et al., 2020 19). We suggest that the lack of evidence supporting trade-offs driving within-species variation does not necessarily mean that physiological costs of reproduction are non-existent (e.g., Smith et al., 1989 and Lemaître et al., 2015 39,46), but rather that, within the wild and within the natural range of reproductive activities, such costs are not relevant to fitness. One key explanation for this effect supported by our meta-analysis (Figure 5) and prior theory 5, is that individuals differ in quality.

Methods

All data and code used in this study can be found at: https://doi.org/10.5061/dryad.q83bk3jnk.

Study sourcing & inclusion criteria

We extracted studies of parental survival to the following year given clutch size raised using the following inclusion criteria (similar to Santos & Nakagawa, 2012 18): the study must be on a wild population; the study must contain variation in the number of offspring produced/raised (hereafter referred to as clutch size for simplicity), the study must report variation in clutch size in relation to parental survival to the following year (including both experimental and observational studies) and must provide sample sizes. We did include studies where parental survival was reported for both parents combined as opposed to Santos & Nakagawa (2012) who required male and female data to be reported separately. Excluded studies and the grounds for their removal are given in the supplementary information (Table S4). We started by, first, extracting data (clutch size raised after manipulation and associated parental survival to the next year) from the included brood-manipulation studies from Santos & Nakagawa (2012) and then searched the literature to include more recently-published studies (Supplementary Methods), and in this search also added studies Santos & Nakagawa missed. In addition to brood manipulation studies, we extracted data from studies that correlated variation in parental survival with natural variation in clutch size (observational studies). For each species used in the brood manipulation studies, we aimed to find the same species for the observational studies to ensure that the effects of quality were estimated across a similar range of species. The reason for this is to ensure the experimental and observational datasets are comparable in terms of species, which could bias results if not considered, and so facilitate a more direct comparison between the quality effect (observational studies) and the trade-off (experimental studies). Where there was no equivalent study in the same species, we attempted to find a study of a congener. In most cases, observational data were obtained from either the same paper as the one describing brood manipulations (11 studies), or via searching for other papers by the same authors (2 studies). If this failed to produce observational data, a search was conducted following the same protocol as for the brood-manipulation experiments, but also specifying species, genus and/or common name in the search (7 studies). Of the 28 species used in the brood manipulation studies, we were able to find 10 of the same species in observational studies. We were able to find a further 6 species which were congeners for observational studies.

From the literature search, 78 individual effect sizes from 30 species and 46 papers were used (20 observational and 58 experimental studies). While extracting these studies we also made note of the average clutch size of the species and the within-species standard deviation in clutch size. We extracted this information from the paper containing the study but if the information was missing, we searched other published literature with the aim to find the information from a similar population (i.e., at a similar latitude).

Extracting effect sizes

We used the absolute value of offspring in the nest (for experimental studies this is the number of offspring after the manipulation occurred) and associated parental survival to estimate an effect size by performing a logistic regression to obtain the log odds ratio for parental survival, given the clutch size (i.e., positive effect sizes indicate an increased chance of survival). For example, a bird who laid a clutch size of 5 but was manipulated to have -1 chick was recorded to have a clutch size of 4. If the manipulation was reported but absolute value of offspring produced was not, we used clutch size to be the species average +/- the number of offspring added or removed. Parental survival was modelled as a binary response variable for the number of birds who survived and who died after raising a given clutch size. Clutch size was averaged (mean) if a single estimate of survival was reported for multiple clutch sizes. ‘Year’ was included as an explanatory variable to correct for between-year variation in adult survival, where data were presented for multiple years. We standardised the clutch size (by the mean of the species and by the within-species standard deviation in clutch size) and transformed clutch size to a proportion of the species mean. For species that have no within-species variance in clutch size, we used a value of 0.01 for the standard deviation in clutch size to prevent issues in calculations when using zero. We, therefore, expressed variation in clutch size in three ways: a raw increase in clutch size, a standardised clutch size and a proportional clutch size.

Meta-analysis

Earlier version: We ran a single model using the log odds ratio calculated for each clutch size transformation (i.e., raw, standardised and proportional) to determine the effect of parental survival, given an increase in reproductive effort using the metafor package 43 in R 3.3.2 44. From these models we were also able to directly compare the effect size of brood manipulation studies and observational studies by including a categorical fixed effect for study type (i.e., experimental or observational). We included phylogeny as a correlation matrix in these meta-analytic models to correct for shared ancestry. The phylogeny was obtained using BEAST to measure a distribution of 1,000 possible phylogenetic trees of the focal 30 species extracted from BirdTree 45. We also included species and each studies’ journal reference as random effects in the model. From these models, we calculated the proportion of variance explained by the phylogenetic effect 46.

Revised version: We ran a single model using the log odds ratio calculated for each clutch size transformation (i.e., raw, standardised and proportional) to determine the effect of parental survival, given an increase in reproductive effort using the metafor package 47 in R 3.3.2 48. From these models we were also able to directly compare the effect size of brood manipulation studies and observational studies by including a categorical fixed effect for study type (i.e., experimental or observational). We included phylogeny as a correlation matrix in these meta-analytic models to correct for shared ancestry. The phylogeny was obtained using BEAST to measure a distribution of 1,000 possible phylogenetic trees of the focal 30 species extracted from BirdTree 49. We also included species and each studies’ journal reference as random effects in the model. From these models, we calculated the proportion of variance explained by the phylogenetic effect 50.

We then tested the effect of the species’ mean clutch size on the relationship between parental survival and clutch size. We ran a single model with the mean centred (from all species used in the meta-analysis) species’ average clutch size in interaction with treatment (brood manipulation or observational). Species, phylogeny and reference were also included as random effects to correct for the similarity of effect sizes within species and studies.

The difference in survival for the different sexes was modelled for each clutch size measure. Brood manipulation studies and observational studies were analysed in separate models. Sex was modelled as a categorical moderator (41 female studies, 27 male studies and 10 mixed studies). Species, phylogeny and reference were included as random effects (Table S2 and Figure S2).

Publication bias

Much of the data used in this analysis were taken from studies where these data were not the main focus of the study. This reduces the risk that our results are heavily influenced by a publication bias for positive results. A funnel plot for the survival against raw clutch size model is presented in Supplementary Information (Figure S3).

Fitness projections

We calculated various isoclines using the brood-manipulation overall effect size (-0.05, based on raw clutch size) that we estimated. Here, an isocline is a trendline representing the change in fitness returns over an individual’s lifespan, given an increase in individual clutch size. An estimated “lifetime reproductive output” was calculated for hypothetical control populations (starting with 100 individuals), where all individuals consistently reproduce at the level of a species mean and have a consistent annual survival rate. We calculated this “lifetime reproductive output” (see Supplementary Information for equation details for calculating lifetime reproductive output) using combinations of a range of species average clutch sizes at 2, 4, 6, 8 and 10 and survival rates of 0.2, 0.3, 0.4, 0.5, 0.6 and 0.7, which reflected the range of clutch sizes and survival rates seen in the species in our meta-analysis. We started with a hypothetical population size of 100 and calculated the populations lifetime reproductive output rate that depends on how many offspring are produced in a clutch and how long individuals in the population live for. We could then use to model the effects of the costs of reproduction using the meta-analytic estimate of how brood size is associated with parental survival (through experimental studies).

The “lifetime reproductive output” estimate was then repeated for a hypothetical population that reproduces at an increased level compared to control, i.e., brood size enlargement, throughout their lives. These individuals therefore produce more offspring but face reductions to survival. This analysis allows us to determine if an overall fitness gain is achieved by producing more offspring despite paying an increased survival cost or whether the survival cost balances out the fitness gains of producing more offspring through reduced reproductive attempts. To obtain this, we added a range of 1–5 offspring to the clutch sizes of the control populations. Using a range of increased clutch sizes allowed us to investigate how increased reproductive effort would affect lifetime fitness. The survival costs were determined by the overall effect size found for brood manipulation studies (per egg). We modelled effect sizes of -0.05, -0.15 and -0.25, which represent, respectively, the meta-analytic overall effect size, its upper confidence interval and a further severe effect within the observed effect sizes (rounded to the closest 0.05 for simplicity). For example, an individual who has an additional offspring in its nest would see a 5%, 14% and 22% (respectively for each effect size) reduction in its survival odds compared to if it reproduced at its normal rate (i.e., the control population rate).

We then calculated the selection differential (LRSbrood manipulation / LRScontrol) between the hypothetical control and “brood manipulation” populations for each combination of survival rate, clutch size and effect size, and plotted this as an isocline. We further plotted the fitness consequences for five exemplar species, where survival rate and clutch size combinations are observable in the wild. We used effect sizes from model predictions at these survival rate and clutch size combinations rather than the meta-analytic mean, thereby providing a biological context.